# Bi621 - In class assignment 7
## Due before class Monday, July 10

The goal of this assignment is to use a dictionary to count the number of times each barcode appears in a sequencing dataset.

Decompress s_1_sequence.txt.gz into your working directory on your own computer. Remember to start with a subset of the data to test your code. (Suggestion: `head -40 s_1_sequence.txt > test.txt` in your terminal) Be sure to change back to s_1_sequence.txt for your final submission.

<!-- grade_id: create_dict, points: 0, grade: false, solution: true -->
```python
barcodes = {}
```

# <font color=red> The next cell contains tests for the autograder. Do not change!

<!-- grade_id: test_create_dict, points: 1, grade: true, solution: false -->
```python
assert barcodes == {}
```

# Populate your dictionary
* Read each FASTQ record and pull the barcode off the sequence from each record.
> [Hint](hint1.ipynb)
* Each barcode will be a key for your dictionary. Increment the counter for each barcode in the dictionary when you encounter that barcode in a sequence.
> [Hint](hint2.txt)

*Remember to test your code each step of the way.*

<!-- grade_id: tally, points: 0, grade: false, solution: true -->
```python
my_file = open("s_1_sequence.txt")
#my_file = open("test_file.txt")
i = 2
for lines in my_file:
    i = i+1
    if i%4 == 0:
        current_code=lines[0:5]
        if barcodes.get(current_code):
            barcodes[current_code] += 1
        else:
            barcodes[current_code]=1
print(barcodes)
my_file.close()        
```

Output:
```
{'CGGCG': 14409, 'ACCAT': 18352, 'CTAGG': 31136, 'CTGAA': 11871, 'AATTT': 23012, 'GAGAT': 11536, 'GAAGC': 18226, 'TGACC': 10931, 'CGATA': 19501, 'TCAGA': 7900, 'TGGTT': 14508, 'GCATT': 26336, 'TCGAG': 18375, 'CTAAT': 75, 'ACTGC': 10659, 'TGGCT': 53, 'ACTTC': 31, 'CAAAG': 15, 'ACTAC': 39, 'TGACT': 190, 'CGGTG': 34, 'CTCAA': 26, 'CTGCT': 52, 'TATTT': 23, 'TCATT': 24, 'TGTCC': 23, 'GAGTT': 90, 'TCGGG': 16, 'CTAGA': 207, 'TTAGG': 19, 'CCTGC': 17, 'CAGGT': 49, 'CTTTG': 13, 'TCAAC': 15, 'TAGAA': 3, 'TCAGT': 80, 'GCCAT': 32, 'CTGAC': 146, 'CCGGT': 10, 'CAGCC': 16, 'TAGCA': 36, 'GCAGC': 25, 'GTTTT': 26, 'CGATG': 47, 'AAGCT': 22, 'AATCT': 32, 'GAAGG': 21, 'GAAAC': 60, 'ACCTC': 12, 'CGGCA': 69, 'TGAAT': 38, 'CCGGG': 4, 'TCCAG': 20, 'CAGAT': 49, 'GAATG': 28, 'CTAAA': 114, 'TAAAA': 14, 'CTAGT': 286, 'CTTAA': 47, 'CGAGA': 35, 'CTGCA': 63, 'CTGGA': 51, 'TCTTT': 17, 'AAAAC': 8, 'AGGTG': 74, 'GATGA': 6, 'TAGGT': 53, 'TTGAT': 6, 'TTCTT': 2, 'TGAAC': 25, 'GAAGA': 58, 'CAATG': 12, 'CTGTT': 47, 'CTGAT': 1
```

# <font color=red> The next cell contains tests for the autograder. Do not change!

<!-- grade_id: test_tally, points: 5, grade: true, solution: false -->
```python
assert len(barcodes) == 976, "Dictionary wrong length"
assert barcodes["ACTGC"] == 10659, "ACTGC count incorrect"
assert barcodes["TTTTT"] == 7, "TTTTT count incorrect"
assert barcodes["CTAGG"] == 31136, "CTAGG count incorrect"
assert barcodes["CTGAG"] == 104, "CTGAG count incorrect"
```

# Print your results
* Print all the barcodes in your dictionary along with their counts. You should order the barcodes from highest counts to fewest counts. The format should be `barcode` [tab] `count`.

<!-- grade_id: Printed_results, points: 1, grade: true, solution: true -->
```python
for key, value in sorted(barcodes.items(), key=lambda x: x[1],reverse=True):
    print (key,"\t",value)
#The code below was my terrible previous attempt!
#bar_list = []
#i=-1
#for each in barcodes.keys():
#    i=i+1
#    count = barcodes.get(each)
#    bar_list.append(each)
#    bar_list[i]=count
#print(sort.bar_list)
```

Output:
```
CTAGG 	 31136
GCATT 	 26336
AATTT 	 23012
CGATA 	 19501
TCGAG 	 18375
ACCAT 	 18352
GAAGC 	 18226
TGGTT 	 14508
CGGCG 	 14409
CTGAA 	 11871
GAGAT 	 11536
TGACC 	 10931
ACTGC 	 10659
TCAGA 	 7900
CTAGT 	 286
CTAGA 	 207
CTAGC 	 203
TGACT 	 190
CTGAT 	 180
CTGAC 	 146
AGATC 	 129
CTAAG 	 118
CTAAA 	 114
CTGAG 	 104
ATTTT 	 99
GAAGT 	 99
CGATT 	 93
GATTT 	 91
TGTTT 	 91
GAGTT 	 90
TCGAT 	 88
CTATG 	 86
GGTTT 	 86
AATTC 	 81
TCAGT 	 80
GAGCT 	 80
TGGTC 	 79
ACCAC 	 78
CGTAT 	 78
CGGCT 	 76
CTGGT 	 76
CTAAT 	 75
AGGTG 	 74
TGACA 	 73
ACATT 	 72
GCATC 	 71
CTGGG 	 70
CGGCA 	 69
CATTT 	 66
ACTGT 	 65
AATTA 	 64
CTGTG 	 64
CTGCA 	 63
GAAAC 	 60
CTGTA 	 60
GAAGA 	 58
GCTTT 	 58
TGGTG 	 57
GAACT 	 56
CGAGT 	 55
TGGTA 	 54
TGGCT 	 53
TAGGT 	 53
ACTGA 	 53
CTGCT 	 52
CCAGG 	 52
CTGGA 	 51
GAATT 	 51
AATTG 	 51
ACCAA 	 50
GCATA 	 50
TGGAG 	 50
CAGGT 	 49
CAGAT 	 49
TCGAC 	 49
TAGAT 	 48
CGATG 	 47
CTTAA 	 47
CTGTT 	 47
ACCAG 	 47
ATGCT 	 46
GGCGT 	 46
TGATT 	 45
GAGAC 	 44
ATAGG 	 43
GAACC 	 43
CGA
```

**CHALLENGE**: Determine how many reads fall below a specified quality score.
* Translate the quality scores for each FASTQ record and determine the average quality score for that read (hint: what encoding?).
* Tally the number of reads that fall below the threshold (you choose the threshold).
* Print the number of sequences read, and the number and the percent of reads that fell below your chosen threshold.

```python

```